# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD28"
COLLABORATORS = "Radosław Matuszczyk 405124, Piotr Łukasiewicz 406163"
```

---

## Zestaw zadań 2: Elementy statystyki opisowej 2 (Notatnik 2/4)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import dataframe_image as dfi
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Wczytanie danych
Wczytaj dane z przygotowanego wcześniej pliku recession_data_1.csv do tablicy `recession_data_1` w taki sposób, aby kolumna Country Code stanowiła indeks tablicy.

Wyświetl przygotowaną tablicę.

<!-- grade_id: cell-184ec13f50fb3bfe, points: 0, grade: true, solution: true -->
```python
recession_data_1 = pd.read_csv('recession_data_1.csv', index_col=0)
recession_data_1
```

### Zadnie 2: Prezentacja danych w formie tabelarycznej [10 pkt]

Przygotuj trzy tablice szeregów rodzielczych przedziałowych częstości wskaźnika GDP growth (annual %) dla poszczególnych lat.

Pierwsza tablica (`recession_table_1`) powinna zawierać szeregi skonstruowane w obarciu o całość danych.

Indeks:
 - Year (wszystkie lata)
 
Kolumny:
 - Kolejne przedziały (-25, -20]	(-20, -15] ... (35, 40] zdefiniowane jako IntervalIndex

Druga tablica (`recession_table_2`) powinna zawierać szeregi skonstruowane indywidualnie dla każdego z rejonów geograficznych.

Indeks:
 - 1 poziomu: Region Name
 - 2 poziomu: Year (wszystkie lata dla każdego z rejonów)
 
Kolumny:
 - Kolejne przedziały (-25, -20], (-20, -15] ... (35, 40] zdefiniowane jako IntervalIndex

Trzecia tablica (`recession_table_3`) powinna zawierać szeregi skonstruowane indywidualnie dla każdego stopnia zamożności.

Indeks:
 - 1 poziomu: Income Group
 - 2 poziomu: Year (wszystkie lata dla każdej z grup)
 
Kolumny:
 - Kolejne przedziały (-25, -20], (-20, -15] ... (35, 40] zdefiniowane jako IntervalIndex
 
Zapisz tablicę `recession_table_1` do pliku recession_table_1.pkl.

Wyświetl przygotowane tablice w taki sposób aby zawarte w tablicy liczby (za wyjątkiem lat i granic przedziałów w nagłówkach kolumn i wierszy) wyświetlały się z dokładnością do 2 miejsca po przecinku.

<!-- grade_id: cell-0b76c13943ac3c6b, points: 0, grade: false, solution: true -->
```python
# Całość danych
def constructSeries(df):
    columns = df.columns[3:]
    bins = pd.IntervalIndex.from_breaks(np.arange(-25, 41, 5))
    length = df.shape[0]
    valueCounts = np.array([pd.cut(df[column], bins).value_counts().sort_index().values / length for column in columns])
    newDf = pd.DataFrame(data=valueCounts, index=columns, columns=bins)
    newDf.index.name = 'Year'
    return newDf

recession_table_1 = constructSeries(recession_data_1)
pd.io.formats.style.Styler(recession_table_1, precision=2)
```

<!-- grade_id: cell-9a1b92247ccba7a2, points: 0, grade: false, solution: true -->
```python
# Podział na rejony
recession_table_2 = recession_data_1.groupby('Region Name').apply(constructSeries)
pd.io.formats.style.Styler(recession_table_2, precision=2)
```

<!-- grade_id: cell-a931928e53b6f707, points: 0, grade: false, solution: true -->
```python
# Podział pod względem stopnia zamożności
recession_table_3 = recession_data_1.groupby('Income Group').apply(constructSeries)
pd.io.formats.style.Styler(recession_table_3, precision=2)
```

<!-- grade_id: cell-c759195a85d2912b, points: 10, grade: true, solution: false -->
```python
### Komórka testowa
## Tablica 1
# Kolumny
#if recession_table_1.columns.inferred_type == 'interval':
lower_boundaries = np.arange(-25, 40, 5)
for i in range(len(lower_boundaries)):
    assert recession_table_1.columns[i] == pd.Interval(lower_boundaries[i], lower_boundaries[i]+5, closed='right')
# Indeks
assert len(set(list(recession_table_1.index)) & set([2006, 2007, 2008, 2009, 2010, 2011])) == 6 \
    or len(set(list(recession_table_1.index)) & set(['2006', '2007', '2008', '2009', '2010', '2011'])) == 6
# Wartości w obrębie tablicy
assert recession_table_1.values.min() >= 0.0
assert recession_table_1.values.max() <= 1.0
assert np.all(np.isclose(recession_table_1.sum(axis=1), 1.0))==True

## Tablica 2
# Kolumny
lower_boundaries = np.arange(-25, 40, 5)
for i in range(len(lower_boundaries)):
    assert recession_table_2.columns[i] == pd.Interval(lower_boundaries[i], lower_boundaries[i]+5, closed='right')
# Indeks
assert set(list(recession_table_2.index.levels[0])) == set(['Africa', 'Americas', 'Asia', 'Europe', 'Oceania'])
set(list(recession_table_2.index.levels[1]))
assert set(list(recession_table_2.index.levels[1])) == set([2006, 2007, 2008, 2009, 2010, 2011]) \
    or set(list(recession_table_2.index.levels[1])) == set(['2006', '2007', '2008', '2009', '2010', '2011'])
# # Wartości w obrębie tablicy
assert recession_table_2.values.min() >= 0.0
assert recession_table_2.values.max() <= 1.0
assert np.all(np.isclose(recession_table_2.sum(axis=1), 1.0))==True

## Tablica 3
# Kolumny
lower_boundaries = np.arange(-25, 40, 5)
for i in range(len(lower_boundaries)):
    assert recession_table_3.columns[i] == pd.Interval(lower_boundaries[i], lower_boundaries[i]+5, closed='right')
# Indeks
assert set(list(recession_table_3.index.levels[0])) == set(['High', 'Low', 'Lower middle', 'Upper middle'])
assert set(list(recession_table_3.index.levels[1])) == set([2006, 2007, 2008, 2009, 2010, 2011]) \
    or set(list(recession_table_3.index.levels[1])) == set(['2006', '2007', '2008', '2009', '2010', '2011'])
# Wartości w obrębie tablicy
assert recession_table_3.values.min() >= 0.0
assert recession_table_3.values.max() <= 1.0
assert np.all(np.isclose(recession_table_3.sum(axis=1), 1.0))==True
```

### Zadanie 3: Prezentacja danych w formie graficznej [25 pkt]

#### a) Diagramy częstości

Bazując na wczytanych do notatnika danych przygotuj diagramy częstości odpowiadające danym zawartym w przygotowanych wcześniej tabelach (podział na identyczne przedziały, skonstruowane w obarciu o całość danych oraz indywidualnie dla każdego z rejonów geograficznych i dla każdego stopnia zamożności, osobne diagramy dla poszczególnych lat). 

Zestaw wyniki na pojedynczej figurze (siatka wykresów 2x5). W górnym wierszu zawszyj diagram skonstruowany w oparciu o całość danych oraz cztery diagramy skonstruowane indywidualnie dla poszczególnych stopni zamożności, a w dolnym wierszu pięć dziagramów skonstruowanych indywidualnie dla poszczególnych rejonów geograficznych). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw identyczne zakresy wartości na odpowiadających sobie osiach wszystkich wykresów,
- ustaw różne od domyślnych i od siebie nawzajem kolory diagramów poszczególnych lat (kolory przyporządkowane poszczególnym latom powinny być identyczne na wszystkich wykresach),
- dodaj do figury lub wykresów legendę zawierającą nazwę wizualizowanego parametru oraz kolory przypisane poszczególnym latom w taki sposób, aby nie zasłaniała/zasłaniały znajdujących się na wykresach danych.

Uwaga! Zadanie należy wykonać w oparciu o wczytane do notatnika szeregi szczegółowe, a nie w oparciu o szeregi rozdzielcze częstości przygotowane w zadaniu 2.

<!-- grade_id: cell-20f667f95ac4349e, points: 10, grade: true, solution: true -->
```python
# Bez uwzględnienia powyższej uwagi...

# fig, ax = plt.subplots(2, 5, figsize=(20,10))
# fig.subplots_adjust(wspace=0.25, hspace=0.5)

# def displayPlot(ax, data, title):
#     data = data.transpose().reset_index()
#     data['index'] = data['index'].astype(str)
#     data.plot(x='index', y=data.columns[1:], kind="bar", ax=ax)
#     ax.set_xlabel('GDP growth (annual %)')
#     ax.set_ylabel('rf')
#     ax.set_title(title)
#     ax.set_ylim([0,0.8])

# displayPlot(ax[0][0], recession_table_1, "All data")
# for ix, income in enumerate(recession_table_3.index.levels[0]):
#     displayPlot(ax[0][ix+1], recession_table_3.loc[income], 'Income = ' + income)
# for ix, region in enumerate(recession_table_2.index.levels[0]):
#     displayPlot(ax[1][ix], recession_table_2.loc[region], 'Region = ' + region)
```

```python
fig, ax = plt.subplots(2, 5, figsize=(20,10))
fig.subplots_adjust(wspace=0.25, hspace=0.5)

def displayPlot(ax, data, title):
    data = data.loc[:,'2006':'2011']
    breaks = np.arange(-25, 41, 5)
    bins = pd.IntervalIndex.from_breaks(breaks).astype(str)
    ax.hist(data, bins=breaks, label=data.columns, weights=np.ones_like(data) / len(data))
    ax.legend()
    labelsLoc = (breaks[1:] + breaks[:-1])/2
    ax.set_xticks(ticks=labelsLoc, labels=bins, rotation=90)
    ax.set_xlabel('GDP growth (annual %)')
    ax.set_ylabel('rf')
    ax.set_title(title)
    ax.set_ylim([0,0.8])

displayPlot(ax[0][0], recession_data_1, "All data")
for ix, income in enumerate(recession_data_1['Income Group'].unique()):
    displayPlot(ax[0][ix+1], recession_data_1[recession_data_1['Income Group'] == income], 'Income = ' + income)
for ix, region in enumerate(recession_data_1['Region Name'].unique()):
    displayPlot(ax[1][ix], recession_data_1[recession_data_1['Region Name'] == region], 'Region = ' + region)
```

#### b) Inne typy wykresów
Bazując na wczytanych do notatnika danych przygotuj wykresy typu boxplot, violinplot, swarmplot i stripplot.

W przypadku każdego z typów wykresów zestaw wyniki na pojedynczej figurze (siatka wykresów 3x1). W górnym wierszu zawszyj wykres skonstruowany w oparciu o całość danych, w wierszu środkowym wykres zawierajacy podział na poszczególne stopnie zamożności, a w wierszu dolnym wykres zawierajacy podział na poszczególne rejony geograficzne. Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw identyczne zakresy wartości na odpowiadających sobie osiach wszystkich wykresów,
- na górnych wykresach dla całości danych użyj pojedynczego koloru jako wypełnienia,
- dla stopni zamożności ustaw jako wypełnienie skalę kolorów coolwarm,
- poszczególnym rejonom geograficznym przypisz jasne/wypłowiałe wersje kolorów przypisanych im na fladze olimpijskiej (Afryka szary, Europa jasnoniebieski itd.),
- w przypadku wykresów swarmplot i stripplot dobierz ich parametry tak by chmury punktów powiązanych z opisanymi w legendzie wariantami parametrów były rozdzielone,
- do dwóch dolnych wykresów dodaj legendę zawierającą nazwę wizualizowanego parametru oraz kolory przypisane poszczególnym rejonom geograficznym/stopniom zamożności w taki sposób, aby nie zasłaniała znajdujących się na wykresach danych.

Uwaga! Zadanie należy wykonać w oparciu o wczytane do notatnika szeregi szczegółowe, a nie w oparciu o szeregi rozdzielcze częstości przygotowane w zadaniu 2.

<!-- grade_id: cell-ae4385a252fa931d, points: 2.5, grade: true, solution: true -->
```python
# Boxplot
fig, ax = plt.subplots(3, 1, figsize=(5,15))
fig.subplots_adjust(hspace=0.25)

def displayPlot(ax, title):
    ax.set_xlabel('Year')
    ax.set_ylabel('GDP growth (annual %)')
    ax.set_title(title)

displayPlot(ax[0], "All data")
displayPlot(ax[1], "Split by income")
displayPlot(ax[2], "Split by region")
sns.boxplot(recession_data_1, ax=ax[0], color='tab:blue')
sns.boxplot(recession_data_1, ax=ax[1], hue='Income Group', palette='coolwarm')
```

<!-- grade_id: cell-bf9c7c83e4f80572, points: 2.5, grade: true, solution: true -->
```python
# Violinplot
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-55f5e27c0761428a, points: 2.5, grade: true, solution: true -->
```python
# Swarmplot
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ede6fce5fff8ebdd, points: 2.5, grade: true, solution: true -->
```python
# Stripplot
# YOUR CODE HERE
raise NotImplementedError()
```

#### c) Dyskusja wyników

Zastanów się jakie są mocne i słabe strony poszczególnych typów wykresów wykonanych w tym zadaniu. W jakich przypadkach wykorzystanie konkretnych typów wykresów do wizualizacji może sprawiać problemy?

<!-- grade_id: cell-343f7f2b58b388ba, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE