# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD28"
COLLABORATORS = "Radosław Matuszczyk 405124, Piotr Łukasiewicz 406163"
```

---

## Zestaw zadań 5: Weryfikacja hipotez statystycznych (Notatnik 2/2)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

```python
# Dane do testów
x = sp.stats.norm.rvs(loc=10, scale=1.5, size=120, random_state=7)
y = sp.stats.norm.rvs(loc=0, scale=0.25, size=120, random_state=10) * x + 5
test_data_1 = pd.DataFrame(data=np.vstack([x, y]).T, columns=["x", "y"])
test_data_1.head()
```

```python
# Dane do testów
x = sp.stats.uniform.rvs(loc=0, scale=10, size=90, random_state=7)
y = 0.75*x  + sp.stats.norm.rvs(loc=0, scale=0.75, size=90, random_state=42) + 2
test_data_2 = pd.DataFrame(data=np.vstack([x, y]).T, columns=["x", "y"])
test_data_2.head()
```

```python
# Dane do testów
x = sp.stats.uniform.rvs(loc=0, scale=10, size=70, random_state=17)
y = x  + 1.25*np.sin(x) + sp.stats.norm.rvs(loc=1, scale=1.5, size=70, random_state=5)
test_data_3 = pd.DataFrame(data=np.vstack([x, y]).T, columns=["x", "y"])
test_data_3.head()
```

### Zadanie 4: Weryfikacja hipotezy o wartości współczynnika prostej regresji [10 pkt]

Przygotuj funkcję `regression_coefficient_NHST()`, która będzie weryfikować hipotezę o wartości współczynnika równania regresji liniowej i wyświetlać wynik testu statystycznego zgodnie z wytycznymi zawartymi w treści zadania oraz w docstring funkcji.

Poza zwracaniem zmiennych wymienionych w docsting funkcji, które są potrzebne do przeprowadzenia automatycznych testów funkcji powinna ona generować wykres zawierający:
- funkcję gęstości prawdopodobieństwa rozkładu teoretycznego wykorzystywanego w teście w zakresie od wartości, dla której dystrybuanta tego rozkładu przyjmuje wartość 0.0001 do wartości, dla której dystrybuanta tego rozkładu przyjmuje wartość 0.9999,
- zaznaczony wyróżniającym się kolorem fragmentem pola pod krzywą odpowiadający wyznaczonemu w teście obszarowi krytycznemu,
- wyróżniającą się kolorem pionowę linię wysokości 1/3 maksymalnej wartości funkcji gęstości prawdopodobieństwa wskazującą obliczoną wartość statystyki t.

Tytuł wykresu powinien zawierać w kolejnych liniach informację o:
- postaci hipotezy zerowej,
- postaci hipotezy alternatywnej,
- wartości wyznaczonych statystyk t i t_alpha,
- zależność między wartościami statystyk decydującą o wyniku testu,
- informację o wyniku testu: "Wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha = A" lub "Odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha = A".

Wszystkie występująca w tytule informacja o testowanym współczynniku (a, b) oraz wartości liczbowe powinny być automatycznie uzupełniane przez funkcję i odpowiednio sformatowane (4 miejsca po przecinku w przypadku t i t_alpha, pozostałe wartości z formatowaniem domyślnym).

Przykład komunikatu zawartego w tytule figury:</br>
    `H0: a = 0.3`</br>
    `H1: a ≠ 0.3`</br>
    `t = 1.2140, t_alpha = 2.1448`</br>
    `|t| < t_alpha`</br>
    `Wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha = 0.95`</br>

<!-- grade_id: cell-76bed301339216ae, points: 0, grade: false, solution: true -->
```python
def regression_coefficient_NHST(data, coef, H0, H1, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks oraz kolumny "x" i "y" z wynikami próby losowej.
    coef: str
        "a" lub "b" w zależności od testowanego współczynnika równania regresji.
    H0: float
        Wartość współczynnika przyjęta jako hipoteza zerowa.
    H1: str
        Postać hipotezy alternatywnej, przyjmuje wartości:
        - two-sided: wartość współczynnika jest różna od wartości przyjętej w H0,
        - less: wartość współczynnika jest mniejsza od wartości przyjętej w H0,
        - greater: wartość współczynnika jest większa od wartości przyjętej w H0.
    alpha: float
        Wartość określająca poziom istotności.
    
    Returns
    -------
    t: float
        Wyliczona na podstawie próby losowej wartość statystyki t.
    t_alpha: float
        Wartość statystyki t_alpha.
    H: int
        Wynik testu statystycznego, przyjmuje wartość:
        0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha,
        1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha.
    """
    # YOUR CODE HERE
    
    res = sp.stats.linregress(data['x'], data['y'])
    a = res.slope
    b = res.intercept
    n = len(data)
    
    y_pred = a*data['x'] + b
    
    S2 = np.sum((data['y'] - y_pred)**2) / (n-2)
    
    
    x_start = sp.stats.t.ppf(0.0001, n-2)
    x_stop = sp.stats.t.ppf(0.9999, n-2)
    x = np.linspace(x_start, x_stop, 1000)
    y = sp.stats.t.pdf(x, n-2)
    plt.figure(figsize=(8,3))
    plt.plot(x, y, 'k')
    
    if coef == 'a':
        D = np.sqrt(S2 / np.sum((data['x'] - np.mean(data['x']))**2))
        t = (a - H0) / D
    elif coef == 'b':
        D = np.sqrt(S2*np.sum(data['x']**2) / (n*np.sum((data['x'] - np.mean(data['x']))**2)))
        t = (b - H0) / D
        
    if H1 == 'two-sided':
        sign = '≠'
        inequality_reject = '|t| >= t_alpha'
        inequality_keep = '|t| < t_alpha'
        t_alpha = sp.stats.t.ppf(1 - alpha/2, n-2)
        H = int(abs(t) >= t_alpha)
        fill_x1 = np.linspace(x_start, -t_alpha, 500)
        fill_x2 = np.linspace(t_alpha, x_stop, 500)
        fill_y1 = sp.stats.t.pdf(fill_x1, n-2)
        fill_y2 = sp.stats.t.pdf(fill_x2, n-2)
        plt.fill_between(fill_x1, fill_y1, color='tab:green')
        plt.fill_between(fill_x2, fill_y2, color='tab:green')
        
    elif H1 == 'less':
        sign = '<'
        inequality_reject = 't <= t_alpha'
        inequality_keep = 't > t_alpha'
        t_alpha = sp.stats.t.ppf(alpha, n-2)
        H = int(t <= t_alpha)
        fill_x = np.linspace(x_start, t_alpha, 500)
        fill_y = sp.stats.t.pdf(fill_x, n-2)
        plt.fill_between(fill_x, fill_y, color='tab:green')
    elif H1 == 'greater':
        sign = '>'
        inequality_reject = 't >= t_alpha'
        inequality_keep = 't < t_alpha'
        t_alpha = sp.stats.t.ppf(1 - alpha, n-2)
        H = int(t >= t_alpha)
        fill_x = np.linspace(t_alpha, x_stop, 500)
        fill_y = sp.stats.t.pdf(fill_x, n-2)
        plt.fill_between(fill_x, fill_y, color='tab:green')
        
    
    plt.vlines(t, ymin=0, ymax=y.max()/3, color='r')
    
    if H:
        plt.title(f"""H0: {coef} = {H0}
H1: {coef} {sign} {H0}
{t = :.4f} {t_alpha = :.4f}
{inequality_reject}
Odrzucenie H0 na rzecz H1 na poziomie istotności {1-alpha = }""")
    else:
        plt.title(f"""H0: {coef} = {H0}
H1: {coef} {sign} {H0}
{t = :.4f} {t_alpha = :.4f}
{inequality_keep}
Wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności {1-alpha = }""")

    return t, t_alpha, H
    
```

<!-- grade_id: cell-a0701afd37928043, points: 10, grade: true, solution: false -->
```python
### Komórka testowa

assert np.all(np.isclose(regression_coefficient_NHST(test_data_1, "a", 0.4, "two-sided", 0.05), (-2.4099957692531735, 1.9802722492407059, 1)))
assert np.all(np.isclose(regression_coefficient_NHST(test_data_1, "b", 4.5, "two-sided", 0.05), (0.18283553669953134, 1.9802722492407059, 0)))
assert np.all(np.isclose(regression_coefficient_NHST(test_data_2, "a", 0.7, "greater", 0.02), (1.6790640072789238, 2.084638091603984, 0)))
assert np.all(np.isclose(regression_coefficient_NHST(test_data_2, "b", 2.4, "less", 0.02), (-2.8374847222952546, -2.084638091603984, 1)))
assert np.all(np.isclose(regression_coefficient_NHST(test_data_2, "a", 0.85, "less", 0.01), (-3.56782781135344, -2.369472273998314, 1)))
assert np.all(np.isclose(regression_coefficient_NHST(test_data_2, "b", 1.5, "greater", 0.01), (2.667124457352881, 2.3694722739983134, 1)))
```

### Zadanie 5: Weryfikacja hipotezy o losowości próby (test serii) [10 pkt]

Przygotuj funkcję `sample_randomness_NHST()`, która będzie weryfikować hipotezę o losowości próby i wyświetlać wynik testu statystycznego zgodnie z wytycznymi zawartymi w treści zadania oraz w docstring funkcji.

Poza zwracaniem zmiennych wymienionych w docsting funkcji, które są potrzebne do przeprowadzenia automatycznych testów funkcji powinna ona generować wykres zawierający:
- funkcję gęstości prawdopodobieństwa rozkładu teoretycznego wykorzystywanego w teście w zakresie od wartości, dla której dystrybuanta tego rozkładu przyjmuje wartość 0.0001 do wartości, dla której dystrybuanta tego rozkładu przyjmuje wartość 0.9999,
- zaznaczony wyróżniającym się kolorem fragmentem pola pod krzywą odpowiadający wyznaczonemu w teście obszarowi krytycznemu,
- wyróżniającą się kolorem pionowę linię wysokości 1/3 maksymalnej wartości funkcji gęstości prawdopodobieństwa wskazującą obliczoną wartość statystyki u.

Tytuł wykresu powinien zawierać w kolejnych liniach informację o:
- postaci hipotezy zerowej,
- postaci hipotezy alternatywnej,
- wartości wyznaczonych statystyk u i u_alpha,
- zależność między wartościami statystyk decydującą o wyniku testu,
- informację o wyniku testu: "Wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha = A" lub "Odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha = A".

Wszystkie występujące w tytule wartości liczbowe powinny być automatycznie uzupełniane przez funkcję i odpowiednio sformatowane (4 miejsca po przecinku w przypadku u i u_alpha, pozostałe wartości z formatowaniem domyślnym).

Przykład komunikatu zawartego w tytule figury:</br>
    `H0: Dobór jednostek do próby jest losowy`</br>
    `H1: Dobór jednostek do próby nie jest losowy`</br>
    `u = 1.3462, u_alpha = 1.6449`</br>
    `|u| < u_alpha`</br>
    `Wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha = 0.95`</br>

<!-- grade_id: cell-24bb3a117e63d1a5, points: 0, grade: false, solution: true -->
```python
def sample_randomness_NHST(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks oraz kolumny "x" i "y" z wynikami próby.
    alpha: float
        Wartość określająca poziom istotności.
    
    Returns
    -------
    u: float
        Wyliczona na podstawie próby wartość statystyki u.
    u_alpha: float
        Wartość statystyki u_alpha.
    H: int
        Wynik testu statystycznego, przyjmuje wartość:
        0 - gdy wynik testu istotności nie daje podstaw do odrzucenia H0 na rzecz H1 na poziomie istotności 1-alpha,
        1 - gdy następuje odrzucenie H0 na rzecz H1 na poziomie istotności 1-alpha.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-243eed59d7b4912a, points: 10, grade: true, solution: false -->
```python
### Komórka testowa

assert np.all(np.isclose(sample_randomness_NHST(test_data_2, 0.05), (1.2774271521862057, 1.959963984540054, 0)))
assert np.all(np.isclose(sample_randomness_NHST(test_data_3, 0.05), (-2.2266930916721734, 1.959963984540054, 1)))
assert np.all(np.isclose(sample_randomness_NHST(test_data_3, 0.02), (-2.2266930916721734, 2.3263478740408408, 0)))
assert np.all(np.isclose(sample_randomness_NHST(test_data_3, 0.1), (-2.2266930916721734, 1.6448536269514722, 1)))
```