# Analiza danych w jzyku Python - wiczenia laboratoryjne 2022/2023

Ten notatnik zalicza si do grupy zestaw贸w zada, na podstawie kt贸rych odbywa si zaliczenie wicze i podlega zwrotowi do oceny w ustalonym na zajciach terminie.

Uwagi i wytyczne og贸lne dotyczce uzupeniania i oceny notatnika:
- Podczas wykonywania zada nale偶y korzysta wycznie z pakiet贸w zaimportowanych na pocztku notatnika oraz z pakiet贸w wchodzcych w skad standardowej biblioteki Pythona, kt贸re mo偶na zaimportowa samodzielnie we wskazanej kom贸rce notatnika.
- Swoje rozwizania nale偶y wprowadza wycznie w miejce nastpujcych fragment贸w kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w kom贸rkach oznaczonych hasem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie nale偶y w 偶aden spos贸b modyfikowa pozostaych fragment贸w kodu oraz innych element贸w notatnika, w szczeg贸lnoci dodawa lub usuwa kom贸rek oraz zmienia nazwy pliku.
- Je偶eli zestaw zada wymaga skorzystania z fragment贸w kodu opracowanego w ramach wczeniejszych zestaw贸w zada nale偶y je umieci we wskazanej kom贸rce notatnika.
- Otrzymywane wyniki i odpowiedzi maj by rezultatem wykonania napisanego kodu, odpowiedzi uzupeniane manualnie nie podlegaj ocenie.
- Zadanie nale偶y wykona w taki spos贸b, aby podczas wykonywania kodu nie zostay wywietlone 偶adne ostrze偶enia.
- Zawarte w notatniku automatyczne testy maj charakter pogldowy. Dotycz one wybranych aspekt贸w zada i maj za zadanie wyapa podstawowe bdy. Przejcie wszystkich test贸w nie oznacza, 偶e zadanie jest wykonane w caoci poprawnie i zostanie ocenione na maksymaln liczb punkt贸w.
- Zadania, kt贸re powoduj wywietlenie komunikatu o bdzie przerywajcym wykonywanie kodu nie podlegaj ocenie.

Uwagi i wytyczne og贸lne dotyczce wizualizacji wynik贸w:
- Wszystkie wykresy powinny by wykonane w jednolitym, przejrzystym i czytelnym stylu, posiada odpowiednio dobrane proporcje i zakresy wartoci osi.
- Wykresy oraz ich osie powinny mie nadane tytuy. Je偶eli w obrbie figury znajduje si wicej ni偶 jeden wykres to figura r贸wnie偶 powinna mie nadany tytu. 
- Figury powinny mie ustawione biae to, tak, aby niezale偶nie od ustawie notatnika wszystkie elementy wykres贸w byy dobrze widoczne (domylnie to jest prze藕roczyste co mo偶e powodowa problemy w notatnikach z ustawionym ciemnym tem).
- Rozmiar poziomy figur nie powinien przekracza 20 cali.

Przed odesaniem zestawu zada do oceny prosz uzupeni kom贸rk z danymi autor贸w rozwizania (`NAME` - nazwa zespou, `COLLABORATORS` - imiona, nazwiska i numery indeks贸w czonk贸w zespou) oraz upewni si, 偶e notatnik dziaa zgodnie z oczekiwaniami. W tym celu nale偶y skorzysta z opcji **Restart Kernel and Run All Cells...** dostpnej na g贸rnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD28"
COLLABORATORS = "Radosaw Matuszczyk 405124, Piotr ukasiewicz 406163"
```

---

## Zestaw zada 4: Wybrane teoretyczne rozkady prawdopodobiestwa oraz elementy teorii estymacji (Notatnik 2/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakiet贸w wchodzcych w skad standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakiet贸w
def variance_estimator(data, population_mean="unknown", data_type="raw_scores"):
    N = len(data['x']) if data_type == "raw_scores" else data['f'].sum()
    biased = population_mean != "unknown"
    
    S = variance_sum_estimator(data, population_mean, data_type)
    
    if biased:
        S /= N
    else:
        S /= N-1
    
    return S

def variance_sum_estimator(data, population_mean="unknown", data_type="raw_scores"):
    if population_mean == "unknown":
        population_mean = mean_estimator(data, data_type)
    
    if data_type == "raw_scores":
        return ((data['x'] - population_mean)**2).sum()
    
    if data_type == "simple_frequency_distribution":
        return (data.index - population_mean)**2 @ data['f']
    
    if data_type == "grouped_frequency_distribution":
        return (data.index.mid - population_mean)**2 @ data['f']
    
def mean_estimator(data, data_type="raw_scores"):
    if data_type == "raw_scores":
        return data.values.mean()
    elif data_type == 'simple_frequency_distribution':
        return (data.index @ data['f']) / data['f'].sum()
    else:
        return data.index.mid @ data['f'] / data['f'].sum()
    
```

```python
# Dane do test贸w dla szereg贸w szczeg贸owych
test_data_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["x"])
test_data_1
```

Output:
```
x
0   5.338105
1   4.906813
2   5.006564
3   5.081503
4   4.842215
5   5.000413
6   4.999822
7   4.649055
8   5.203532
9   5.120100
10  4.874914
11  4.965690
12  5.101060
13  4.947729
14  4.951450
15  4.709352
16  5.110916
17  5.024776
18  5.054892
19  4.694695
```

```python
# Dane do test贸w dla szereg贸w rozdzielczych
test_population = pd.DataFrame(data=sp.special.round(sp.stats.norm.rvs(loc=10, scale=3, size=10000, random_state=3)), columns=["x"])
raw_test_data_2 = test_population.sample(50, random_state=21)
test_data_2 = raw_test_data_2.groupby('x').size().reset_index(name='f').rename(columns={'Col1':'Col_value'})
test_data_2.set_index('x', inplace=True)
test_data_2
```

Output:
```
f
x      
4.0   3
5.0   1
6.0   1
7.0   7
8.0   7
9.0   5
10.0  8
11.0  5
12.0  2
13.0  6
14.0  3
15.0  1
18.0  1
```

```python
# Dane do test贸w dla szereg贸w rozdzielczych przedziaowych
raw_test_data_3 = pd.DataFrame(data=sp.stats.norm.rvs(loc=-5, scale=2, size=100, random_state=25), columns=["x"])
bins = np.arange(-11, 2, 1) 
index = pd.Index(pd.cut(raw_test_data_3['x'], bins).cat.categories)
test_data_3 = pd.DataFrame(data=raw_test_data_3.groupby(pd.cut(raw_test_data_3['x'], bins)).count().values, index=index)
test_data_3.index.rename("x", inplace=True)
test_data_3.rename(columns= {0: "f"}, inplace=True)
test_data_3
```

Output:
```
f
x             
(-11, -10]   1
(-10, -9]    2
(-9, -8]     7
(-8, -7]     7
(-7, -6]    20
(-6, -5]    16
(-5, -4]    21
(-4, -3]     8
(-3, -2]     8
(-2, -1]     6
(-1, 0]      3
(0, 1]       1
```

### Zadanie 3: Estymacja przedziaowa wartoci oczekiwanej [15 pkt]

Przygotuj funkcj `interval_mean_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci oczekiwanej dla trzech rodzaj贸w danych wejciowych - szeregu szczeg贸owego, szeregu rozdzielczego i szeregu rozdzielczego przedziaowego.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

` nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-3162b1779d0c603e, points: 0, grade: false, solution: true -->
```python
def interval_mean_estimation(data, alpha, population_std="unknown", data_type="raw_scores"):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca w zale偶noci od wartoci parametru data_type:
        - domylny indeks i pojedyncz kolumn "x" z wartociami pr贸by losowej,
        - indeks o nazwie "x" z wariantami cechy oraz pojedyncz kolumn "f" ze zliczeniami wariant贸w,
        - indeks (IntervalIndex) o nazwie "x" z przedziaami wartoci cechy oraz pojedyncz kolumn "f" ze zliczeniami wariant贸w w poszczeg贸lnych przedziaach.
    alpha: float
        Warto okrelajca poziom istotnoci.
    population_std: str or float
        Je偶eli odchylenie standardowe populacji nie jest znane to parametr przyjmuje warto "unknown",
        w przeciwnym wypadku jest to warto okrelajca odchylenie standardowe populacji.
    data_type: str
        Okrela typ szeregu statystycznego:
        - "raw_scores" - szereg szczeg贸owy,
        - "simple_frequency_distribution" - szereg rozdzielczy,
        - "grouped_frequency_distribution" - szereg rozdzielczy przedziaowy.
    
    Returns
    -------
    mean_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    mean_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    data_mean = mean_estimator(data, data_type)
    n = len(data['x']) if data_type == "raw_scores" else data['f'].sum()
    n_sqrt = np.sqrt(n)
    if population_std != 'unknown':
        u = sp.stats.norm.ppf(1-alpha/2)
        mean_low = data_mean - u*population_std/n_sqrt
        mean_high = data_mean + u*population_std/n_sqrt
    else:
        t = sp.stats.t.ppf(1-alpha/2, n-1)
        S = variance_estimator(data, 'unknown', data_type)
        mean_low = data_mean - t*np.sqrt(S)/n_sqrt
        mean_high = data_mean + t*np.sqrt(S)/n_sqrt
    print(" nale偶y do przedziau [{}, {}] przy zao偶eniu poziomu ufnoci 1- = {}".format(mean_low, mean_high, 1-alpha))
    return [mean_low, mean_high]
```

<!-- grade_id: cell-37f4b7bd74d39411, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Testy dla szereg贸w szczeg贸owych

# Test na test_data_1 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_mean_estimation(test_data_1, 0.1, population_std=0.2, data_type="raw_scores"), (4.905619719495418, 5.05273990041144))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_1, 0.05, population_std=0.2, data_type="raw_scores"), (4.891527555895771, 5.066832064011087))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_1, 0.02, population_std=0.2, data_type="raw_scores"), (4.875142370240081, 5.083217249666777))))
print()

# # Test na test_data_1 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_mean_estimation(test_data_1, 0.1, data_type="raw_scores"), (4.913356369231727, 5.045003250675131))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_1, 0.05, data_type="raw_scores"), (4.899504007229391, 5.058855612677467))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_1, 0.02, data_type="raw_scores"), (4.882508507351398, 5.07585111255546))))
```

Output:
```
nale偶y do przedziau [4.905619719495418, 5.05273990041144] przy zao偶eniu poziomu ufnoci 1- = 0.9

 nale偶y do przedziau [4.891527555895771, 5.066832064011087] przy zao偶eniu poziomu ufnoci 1- = 0.95

 nale偶y do przedziau [4.875142370240081, 5.083217249666777] przy zao偶eniu poziomu ufnoci 1- = 0.98

 nale偶y do przedziau [4.913356369231727, 5.045003250675131] przy zao偶eniu poziomu ufnoci 1- = 0.9

 nale偶y do przedziau [4.899504007229391, 5.058855612677467] przy zao偶eniu poziomu ufnoci 1- = 0.95

 nale偶y do przedziau [4.882508507351398, 5.07585111255546] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

<!-- grade_id: cell-b698fbe7314a1a79, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Testy dla szereg贸w rozdzielczych

# Test na test_data_2 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_mean_estimation(test_data_2, 0.1, population_std=3, data_type="simple_frequency_distribution"), (9.002147707793995, 10.397852292206004))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_2, 0.05, population_std=3, data_type="simple_frequency_distribution"), (8.868457705390192, 10.531542294609807))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_2, 0.02, population_std=3, data_type="simple_frequency_distribution"), (8.713014185720088, 10.686985814279911))))
print()

# Test na test_data_2 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_mean_estimation(test_data_2, 0.1, data_type="simple_frequency_distribution"), (8.985882655267286, 10.414117344732713))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_2, 0.05, data_type="simple_frequency_distribution"), (8.844032825120204, 10.555967174879795))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_2, 0.02, data_type="simple_frequency_distribution"), (8.675649993053778, 10.72435000694622))))
```

Output:
```
nale偶y do przedziau [9.002147707793995, 10.397852292206004] przy zao偶eniu poziomu ufnoci 1- = 0.9

 nale偶y do przedziau [8.868457705390192, 10.531542294609807] przy zao偶eniu poziomu ufnoci 1- = 0.95

 nale偶y do przedziau [8.713014185720088, 10.686985814279911] przy zao偶eniu poziomu ufnoci 1- = 0.98

 nale偶y do przedziau [8.985882655267286, 10.414117344732713] przy zao偶eniu poziomu ufnoci 1- = 0.9

 nale偶y do przedziau [8.844032825120204, 10.555967174879795] przy zao偶eniu poziomu ufnoci 1- = 0.95

 nale偶y do przedziau [8.675649993053778, 10.72435000694622] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

<!-- grade_id: cell-e04d885c1db8d9d7, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Testy dla szereg贸w rozdzielczych przedziaowych

# Test na test_data_3 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_mean_estimation(test_data_3, 0.1, population_std=2, data_type="grouped_frequency_distribution"), (-5.448970725390295, -4.791029274609706))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_3, 0.05, population_std=2, data_type="grouped_frequency_distribution"), (-5.511992796908011, -4.728007203091989))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_3, 0.02, population_std=2, data_type="grouped_frequency_distribution"), (-5.585269574808168, -4.654730425191832))))
print()

# Test na test_data_3 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_mean_estimation(test_data_3, 0.1, data_type="grouped_frequency_distribution"), (-5.487717471820824, -4.752282528179176))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_3, 0.05, data_type="grouped_frequency_distribution"), (-5.5594333457618, -4.6805666542382))))
print()
assert(np.all(np.isclose(interval_mean_estimation(test_data_3, 0.02, data_type="grouped_frequency_distribution"), (-5.643675933878456, -4.596324066121544))))
```

Output:
```
nale偶y do przedziau [-5.448970725390295, -4.791029274609706] przy zao偶eniu poziomu ufnoci 1- = 0.9

 nale偶y do przedziau [-5.511992796908011, -4.728007203091989] przy zao偶eniu poziomu ufnoci 1- = 0.95

 nale偶y do przedziau [-5.585269574808168, -4.654730425191832] przy zao偶eniu poziomu ufnoci 1- = 0.98

 nale偶y do przedziau [-5.487717471820824, -4.752282528179176] przy zao偶eniu poziomu ufnoci 1- = 0.9

 nale偶y do przedziau [-5.5594333457618, -4.6805666542382] przy zao偶eniu poziomu ufnoci 1- = 0.95

 nale偶y do przedziau [-5.643675933878456, -4.596324066121544] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

### Zadanie 4: Estymacja przedziaowa wariancji [15 pkt]
Przygotuj funkcj `interval_variance_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wariancji dla trzech rodzaj贸w danych wejciowych - szeregu szczeg贸owego, szeregu rozdzielczego i szeregu rozdzielczego przedziaowego.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`^2  nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-1f107aa4297b1e13, points: 0, grade: false, solution: true -->
```python
def interval_variance_estimation(data, alpha, population_mean="unknown", data_type="raw_scores"):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca w zale偶noci od wartoci parametru data_type:
        - domylny indeks i pojedyncz kolumn "x" z wartociami pr贸by losowej,
        - indeks o nazwie "x" z wariantami cechy oraz pojedyncz kolumn "f" ze zliczeniami wariant贸w,
        - indeks (IntervalIndex) o nazwie "x" z przedziaami wartoci cechy oraz pojedyncz kolumn "f" ze zliczeniami wariant贸w w poszczeg贸lnych przedziaach.
    alpha: float
        Warto okrelajca poziom istotnoci.
    population_mean: str or float
        Je偶eli warto oczekiwana populacji nie jest znane to parametr przyjmuje warto "unknown",
        w przeciwnym wypadku jest to warto okrelajca warto oczekiwan populacji.
    data_type: str
        Okrela typ szeregu statystycznego:
        - "raw_scores" - szereg szczeg贸owy,
        - "simple_frequency_distribution" - szereg rozdzielczy,
        - "grouped_frequency_distribution" - szereg rozdzielczy przedziaowy.
    
    Returns
    -------
    var_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    var_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    N = len(data['x']) if data_type == "raw_scores" else data['f'].sum()

    S = variance_sum_estimator(data, population_mean, data_type)
    
    low = S / sp.stats.chi2.ppf(1-alpha/2, N-1)
    high = S / sp.stats.chi2.ppf(alpha/2, N-1)
    print(f'^2 nale偶y do przedziau [{low :.4f}, {high :.4f}] przy zao偶eniu poziomu ufnoci 1- = {1 - alpha}')
    return (low, high)
```

<!-- grade_id: cell-bd878d13cb44f436, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Testy dla szereg贸w szczeg贸owych

# Test na test_data_1 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_variance_estimation(test_data_1, 0.1, population_mean=5, data_type="raw_scores"), (0.01855573624556403, 0.05528660848892888))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_1, 0.05, population_mean=5, data_type="raw_scores"), (0.017025745016269588, 0.06280068548350207))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_1, 0.02, population_mean=5, data_type="raw_scores"), (0.015455150810030724, 0.07328116757245808))));
print()

# Test na test_data_1 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_variance_estimation(test_data_1, 0.1, data_type="raw_scores"), (0.018268125369359626, 0.05442967510189219))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_1, 0.05, data_type="raw_scores"), (0.0167618487538223, 0.06182728513230527))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_1, 0.02, data_type="raw_scores"), (0.015215598500840796, 0.07214532146342346))));
```

Output:
```
^2 nale偶y do przedziau [0.0186, 0.0553] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [0.0170, 0.0628] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [0.0155, 0.0733] przy zao偶eniu poziomu ufnoci 1- = 0.98

^2 nale偶y do przedziau [0.0183, 0.0544] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [0.0168, 0.0618] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [0.0152, 0.0721] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

<!-- grade_id: cell-56fe2afee84070b5, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Testy dla szereg贸w rozdzielczych

# Test na test_data_2 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_variance_estimation(test_data_2, 0.1, population_mean=10, data_type="simple_frequency_distribution"), (6.7683018526269425, 13.23300783223187))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_2, 0.05, population_mean=10, data_type="simple_frequency_distribution"), (6.393969919236964, 14.229161421841045))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_2, 0.02, population_mean=10, data_type="simple_frequency_distribution"), (5.993101315036716, 15.514512026199176))));
print()

# Test na test_data_2 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_variance_estimation(test_data_2, 0.1, data_type="simple_frequency_distribution"), (6.700468092411306, 13.100383032131553))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_2, 0.05, data_type="simple_frequency_distribution"), (6.329887815369333, 14.086552899795869))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_2, 0.02, data_type="simple_frequency_distribution"), (5.933036825242363, 15.359021371148184))));
```

Output:
```
^2 nale偶y do przedziau [6.7683, 13.2330] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [6.3940, 14.2292] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [5.9931, 15.5145] przy zao偶eniu poziomu ufnoci 1- = 0.98

^2 nale偶y do przedziau [6.7005, 13.1004] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [6.3299, 14.0866] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [5.9330, 15.3590] przy zao偶eniu poziomu ufnoci 1- = 0.98
```

<!-- grade_id: cell-edeee68cbb07ad2f, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Testy dla szereg贸w rozdzielczych przedziaowych

# Test na test_data_3 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_variance_estimation(test_data_3, 0.1, population_mean=-5, data_type="grouped_frequency_distribution"), (3.9521130029725247, 6.320871976892474))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_3, 0.05, population_mean=-5, data_type="grouped_frequency_distribution"), (3.7921854749549446, 6.63839734543416))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_3, 0.02, population_mean=-5, data_type="grouped_frequency_distribution"), (3.6170094460586584, 7.034533746042385))));
print()

# Test na test_data_3 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(interval_variance_estimation(test_data_3, 0.1, data_type="grouped_frequency_distribution"), (3.940427083620819, 6.302181924229793))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_3, 0.05, data_type="grouped_frequency_distribution"), (3.7809724419283834, 6.618768408724867))));
print()
assert(np.all(np.isclose(interval_variance_estimation(test_data_3, 0.02, data_type="grouped_frequency_distribution"), (3.606314387326986, 7.013733481988378))));
```

Output:
```
^2 nale偶y do przedziau [3.9521, 6.3209] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [3.7922, 6.6384] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [3.6170, 7.0345] przy zao偶eniu poziomu ufnoci 1- = 0.98

^2 nale偶y do przedziau [3.9404, 6.3022] przy zao偶eniu poziomu ufnoci 1- = 0.9

^2 nale偶y do przedziau [3.7810, 6.6188] przy zao偶eniu poziomu ufnoci 1- = 0.95

^2 nale偶y do przedziau [3.6063, 7.0137] przy zao偶eniu poziomu ufnoci 1- = 0.98
```