# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD28"
COLLABORATORS = "Radosław Matuszczyk 405124, Piotr Łukasiewicz 406163"
```

---

## Zestaw zadań 4: Wybrane teoretyczne rozkłady prawdopodobieństwa oraz elementy teorii estymacji (Notatnik 1/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Podstawowe rozkłady teoretyczne wykorzystywane w statystyce matematycznej [15 pkt]
Przeanalizuj własności podstawowych teoretycznych rozkładów prawdopodobieństwa wykorzystywane w statystyce matematycznej - rozkładu normalnego, rozkładu t-Studenta, rozkładu chi-kwadrat i rozkładu F Snedecora.

#### a) Wizualizacja

Zwizualizuj funkcje gęstości prawdopodobieństwa następujących rozkładów teoretycznych:
- Wykres 1: $N(0, 1)$, $N(0, 0.75)$, $N(0, 1.25)$, $N(-2, 1)$, $N(2, 1)$
- Wykres 2: $N(0, 1)$, $t(1)$, $t(2)$, $t(5)$, $t(30)$
- Wykres 3: $\chi^2(1)$, $\chi^2(2)$, $\chi^2(3)$, $\chi^2(5)$, $\chi^2(8)$
- Wykres 4: $F(10, 1)$, $F(10, 2)$, $F(10, 5)$, $F(10, 10)$, $F(10, 20)$
- Wykres 5: $F(1, 10)$, $F(2, 10)$, $F(5, 10)$, $F(10, 10)$, $F(20, 10)$
- Wykres 6: $F(1, 1)$, $F(2, 2)$, $F(5, 5)$, $F(25, 25)$, $F(50, 50)$

gdzie:
- $N(\mu, \sigma)$ - rozkład normalny o wartości oczekiwanej $\mu$ i odchyleniu standardowym $\sigma$,
- $t(df)$ - rozkład t-Studenta o $df$ stopniach swobody,
- $\chi^2(df)$ - rozkład chi-kwadrat o $df$ stopniach swobody,
- $F(df1, df2)$ - rozkład F Snedecora o $df1$ i $df2$ stopniach swobody.

Zestaw wykresy na pojedynczej figurze (siatka wykresów 2x3). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw identyczne zakresy osi x i y wykresów 1, 2,
- ustaw identyczne zakresy osi x i y wykresów 4, 5, 6,
- dobierz zakresy osi x wszystkiech wykresów w taki sposób, aby dobrze widoczny był charakter znajdujących się na nich funkcji,
- zadbaj o to aby rozkład $N(0, 1)$ na obu wykresach był oznaczony w ten sam sposób (kolor, grubość, styl linii itd.),
- do każdego wykresów dodaj legendę opisującą widoczne na nim rozkłady i zadbaj o to, aby legenda nie zasłaniała innych elementów wykresu.

<!-- grade_id: cell-4831d3956629a5fb, points: 10, grade: true, solution: true -->
```python
x1 = np.linspace(-5, 5, 1000)
x2 = np.linspace(0, 10, 1000)

plot1 = [
    sp.stats.norm.pdf(x1, loc=0, scale=1),
    sp.stats.norm.pdf(x1, loc=0, scale=0.75),
    sp.stats.norm.pdf(x1, loc=0, scale=1.25),
    sp.stats.norm.pdf(x1, loc=-2, scale=1),
    sp.stats.norm.pdf(x1, loc=2, scale=1)
]
plot2 = [
    sp.stats.norm.pdf(x1, loc=0, scale=1),
    sp.stats.t.pdf(x1, df=1),
    sp.stats.t.pdf(x1, df=2),
    sp.stats.t.pdf(x1, df=5),
    sp.stats.t.pdf(x1, df=30)
]
plot3 = [
    sp.stats.chi2.pdf(x2, df=1),
    sp.stats.chi2.pdf(x2, df=2),
    sp.stats.chi2.pdf(x2, df=3),
    sp.stats.chi2.pdf(x2, df=5),
    sp.stats.chi2.pdf(x2, df=8)
]
plot4 = [
    sp.stats.f.pdf(x2, dfn=10, dfd=1),
    sp.stats.f.pdf(x2, dfn=10, dfd=2),
    sp.stats.f.pdf(x2, dfn=10, dfd=5),
    sp.stats.f.pdf(x2, dfn=10, dfd=10),
    sp.stats.f.pdf(x2, dfn=10, dfd=20)
]
plot5 = [
    sp.stats.f.pdf(x2, dfn=1, dfd=10),
    sp.stats.f.pdf(x2, dfn=2, dfd=10),
    sp.stats.f.pdf(x2, dfn=5, dfd=10),
    sp.stats.f.pdf(x2, dfn=10, dfd=10),
    sp.stats.f.pdf(x2, dfn=20, dfd=10)
]
plot6 = [
    sp.stats.f.pdf(x2, dfn=1, dfd=1),
    sp.stats.f.pdf(x2, dfn=2, dfd=2),
    sp.stats.f.pdf(x2, dfn=5, dfd=5),
    sp.stats.f.pdf(x2, dfn=25, dfd=25),
    sp.stats.f.pdf(x2, dfn=50, dfd=50)
]

colors = ['black', 'blue', 'green', 'tab:red', 'orange']

fig, axes = plt.subplots(2, 3, figsize=(16,8))
fig.tight_layout(h_pad=4, w_pad=4)
fig.suptitle('Wizualizacja funkcji gęstości prawdopodobieństwa podstawowych rozkładów', fontsize=18, y=1.075)

def display_plot(ax, x, plots, range_x, range_y, title, labels):
    for i, pplot in enumerate(plots):
        ax.plot(x, pplot, color=colors[i])
        ax.set_xlim(range_x)
        ax.set_ylim(range_y)
        ax.set_xlabel('x')
        ax.set_ylabel('y')
        ax.set_title(title)
        ax.legend(handles=[plt.Line2D([], [], color=colors[0], label=labels[0]),
                           plt.Line2D([], [], color=colors[1], label=labels[1]),
                           plt.Line2D([], [], color=colors[2], label=labels[2]),
                           plt.Line2D([], [], color=colors[3], label=labels[3]),
                           plt.Line2D([], [], color=colors[4], label=labels[4])], 
                  loc='best', fontsize=10, title='rozkład')

display_plot(axes[0][0], x1, plot1, [-5, 5], [0, 0.8], title='rozkład normalny', labels=['N(0,1)', 'N(0,0.75)', 'N(0,1.25)', 'N(−2,1)', 'N(2,1)'])
display_plot(axes[0][1], x1, plot2, [-5, 5], [0, 0.8], title='rozkład t-Studenta', labels=['N(0,1)', 't(1)', 't(2)', 't(5)', 't(30)'])
display_plot(axes[0][2], x2, plot3, [0, 10], [0, 0.8], title='rozkład chi2', labels=['chi2(1)', 'chi2(2)', 'chi2(3)', 'chi2(5)', 'chi2(8)'])
display_plot(axes[1][0], x2, plot4, [0, 3], [0, 1.6], title='rozkład F Snedecora', labels=['F(10,1)', 'F(10,2)', 'F(10,5)', 'F(10,10)', 'F(10,20)'])
display_plot(axes[1][1], x2, plot5, [0, 3], [0, 1.6], title='rozkład F Snedecora', labels=['F(1,10)', 'F(2,10)', 'F(5,10)', 'F(10,10)', 'F(20,10)'])
display_plot(axes[1][2], x2, plot6, [0, 3], [0, 1.6], title='rozkład F Snedecora', labels=['F(1,1)', 'F(2,2)', 'F(5,5)', 'F(25,25)', 'F(50,50)'])
```

#### b) Charakterystyka rozkładów
Na podstawie przygotowanych wykresów scharakteryzuje opisane rozkłady teoretyczne i odpowiedz na pytanie w jaki sposób zmiana wartości parametrów rozkładów wpływa na położenie i kształt ich funkcji gęstości prawdopodobieństwa.

<!-- grade_id: cell-d6f6af864a438137, points: 5, grade: true, solution: true -->
**Rozkład normalny** - rozkład na kształt krzywej dzwonowej, jest symetryczny. Średnia wpływa na lokalizację piku, a odchylenie standardowe na szerokość "dzwonu". Im większe odchylenie tym szerszy "dzwon".

**Rozkład t-Studenta** - dla dużych stopni swobody przypomina rozkład normalny, dla małych jest nieco szerszy, tzn. wartości mocno odbiegające od średniej są bardziej prawdopodobne niż w przypadku rozkładu normalnego. Jest rozkładem symetrycznym, jednomodalnym.

**Rozkład chi2** - rozkład nieujemny, jednomodalny, niesymetryczny, prawostronny - wartości gęstości prawdopodobieństwa z prawej strony piku wolniej się zmniejszają. Wraz ze wzrostem stopnia swobody, rozkład staje się coraz bardziej spłaszczony i wydłużony, a pik przesuwa się w prawą stronę.

**Rozkład F Snedecora** - rozkład nieujemny, jednomodalny, niesymetryczny, prawostronny. Zwiększenie stopni swobody licznika (dfn) powoduje, że rozkład staje się bardziej symetryczny - obszar centralny staje się bardziej koncentryczny. Zwiększenie stopni swobody mianownika (dfd) skutkuje zmniejszeniem skośności rozkładu - pik przesuwany jest w kierunku większych wartości, a różnica między prawą i lewą stroną rozkładu maleje. Dla bardzo dużych wartości stopni swobody, rozkład F zbiega do rozkładu normalnego.

### Zadanie 2: Praca z rozkładami teoretycznymi [10 pkt]

#### a) Wyliczanie prawdopodobieństw na podstawie wartości

Oblicz i zapisz do zmiennych o podanych nazwach następujące wartości prawdopodobieństwa:
- `p1` - $P(X < -1.5)$, gdzie $X$ to zmienna losowa ciągła posiadająca rozkład t-Studenta o 10 stopniach swobody - $t(10)$,
- `p2` - $P(X > 2)$, gdzie $X$ to zmienna losowa ciągła posiadająca rozkład F Snedecora o 5 i 50 stopniach swobody - $F(5, 50)$,
- `p3` - $P(3 < X < 10)$, gdzie $X$ to zmienna losowa ciągła posiadająca rozkład chi-kwadrat o 4 stopniach swobody - $\chi^2(4)$,
- `p4` - $P(X < -2.75 | X > 1.5)$, gdzie $X$ to zmienna losowa ciągła posiadająca rozkład normalny o wartości oczewiwanej równej -0.5 i odchyleniu standardowym równym 1.5 - $N(-0.5, 1.5)$,

Każde z rozwiązań zilustruj wykresem zawierającym krzywą gęstości prawdopodobieństwa rozkładu z zaznaczonym kolorem fragmentem pola pod krzywą odpowiadający wyznaczonej wartości prawdopodobieństwa. Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw zakres osi x od wartości zmiennej $X$, dla której dystrybuanta analizowanego rozkładu przyjmuje wartość 0.0001 do wartości zmiennej $X$, dla której dystrybuanta analizowanego rozkładu przyjmuje wartość 0.9999,
- w tytule wykresu zawrzyj uzupełniony wzór na prawdopodobieństwo uzupełniony o wyliczoną wartość prawdopodobieństwa podaną do 4 miejsca po przecinku (np. P(X < -1.5) = x.xxxx, gdzie x.xxxx to obliczona wartość).

<!-- grade_id: cell-e8b63c48ec18e621, points: 0, grade: false, solution: true -->
```python
# p1
p1 = sp.stats.t.cdf(x=-1.5, df=10)

xL = sp.stats.t.ppf(q=0.0001, df=10)
xR = sp.stats.t.ppf(q=0.9999, df=10)
x = np.linspace(xL, xR, 1000)
y = sp.stats.t.pdf(x=x, df=10)
plt.plot(x, y, 'k')
stop = np.searchsorted(x, -1.5)
plt.fill_between(x[:stop], y[:stop], color='tab:green')
plt.title('P(X < -1.5) = {:.4f}'.format(p1))
plt.xlabel('x')
plt.ylabel('y');
```

<!-- grade_id: cell-e1993944ea90eb0f, points: 0, grade: false, solution: true -->
```python
# p2
p2 = 1 - sp.stats.f.cdf(x=2, dfn=5, dfd=50)

xL = sp.stats.f.ppf(q=0.0001, dfn=5, dfd=50)
xR = sp.stats.f.ppf(q=0.9999, dfn=5, dfd=50)
x = np.linspace(xL, xR, 1000)
y = sp.stats.f.pdf(x=x, dfn=5, dfd=50)
plt.plot(x, y, 'k')
start = np.searchsorted(x, 2)
plt.fill_between(x[start:], y[start:], color='tab:green')
plt.title('P(X > 2) = {:.4f}'.format(p2))
plt.xlabel('x')
plt.ylabel('y');
```

<!-- grade_id: cell-8844111764c19f52, points: 0, grade: false, solution: true -->
```python
# p3
p3 = sp.stats.chi2.cdf(x=10, df=4) - sp.stats.chi2.cdf(x=3, df=4)

xL = sp.stats.chi2.ppf(q=0.0001, df=4)
xR = sp.stats.chi2.ppf(q=0.9999, df=4)
x = np.linspace(xL, xR, 1000)
y = sp.stats.chi2.pdf(x=x, df=4)
plt.plot(x, y, 'k')
start = np.searchsorted(x, 3)
stop = np.searchsorted(x, 10)
plt.fill_between(x[start:stop], y[start:stop], color='tab:green')
plt.title('P(3 < X < 10) = {:.4f}'.format(p3))
plt.xlabel('x')
plt.ylabel('y');
```

<!-- grade_id: cell-65717becaa4148fc, points: 0, grade: false, solution: true -->
```python
# p4
p4 = sp.stats.norm.cdf(x=-2.75, loc=-0.5, scale=1.5) + 1 - sp.stats.norm.cdf(x=1.5, loc=-0.5, scale=1.5)

xL = sp.stats.norm.ppf(q=0.0001, loc=-0.5, scale=1.5)
xR = sp.stats.norm.ppf(q=0.9999, loc=-0.5, scale=1.5)
x = np.linspace(xL, xR, 1000)
y = sp.stats.norm.pdf(x=x, loc=-0.5, scale=1.5)
plt.plot(x, y, 'k')
stop = np.searchsorted(x, -2.75)
start = np.searchsorted(x, 1.5)
plt.fill_between(x[:stop], y[:stop], color='tab:green')
plt.fill_between(x[start:], y[start:], color='tab:green')
plt.title('P(X < -2.75 | X > 1.5) = {:.4f}'.format(p4))
plt.xlabel('x')
plt.ylabel('y');
```

<!-- grade_id: cell-76ab68be8461d402, points: 5, grade: true, solution: false -->
```python
### Komórka testowa
assert p1 == 0.08225366322272008
assert p2 == 0.0947529020300999
assert p3 == 0.5173977183765621
assert p4 == 0.15801842099472596
```

#### b) Wyliczanie wartości na podstawie prawdopodobieństw

Rozwiąż i zapisz do zmiennych o podanych nazwach rozwiązania następujących równań:
- `s1` - $P(X < s) = 0.25$, gdzie X to zmienna losowa ciągła posiadająca rozkład chi-kwadrat o 5 stopniach swobody - $\chi^2(5)$,
- `s2` - $P(X > s) = 0.15$, gdzie X to zmienna losowa ciągła posiadająca rozkład F Snedecora o 5 i 20 stopniach swobody - $F(5, 20)$,
- `s3` - $P(-1.25 < X < s) = 0.35$, gdzie x to zmienna losowa ciągła posiadająca rozkład normalny o wartości oczewiwanej równej -0.4 i odchyleniu standardowym równym 0.75 - $N(-0.4, 0.75)$,

Każde z rozwiązań zilustruj wykresem zawierającym krzywą gęstości prawdopodobieństwa rozkładu z zaznaczonym kolorem fragmentem pola pod krzywą odpowiadający wartości prawdopodobieństwa określonej przez rozwiązywane równanie. Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw zakres osi x od wartości zmiennej $X$, dla której dystrybuanta analizowanego rozkładu przyjmuje wartość 0.0001 do wartości zmiennej $X$, dla której dystrybuanta analizowanego rozkładu przyjmuje wartość 0.9999,
- oznacz krawędź obszaru, na której znajduje się wyznaczona wartość $s$ linią w wyróżniającym się kolorze,
- w tytule wykresu zawrzyj uzupełniony wzór na prawdopodobieństwo uzupełniony o wyliczoną wartość zmiennej $X$ podaną do 4 miejsca po przecinku (np. P(X < s.ssss) = 0.25, gdzie s.ssss to obliczona wartość).

<!-- grade_id: cell-7c67b0379d415bec, points: 0, grade: false, solution: true -->
```python
# s1
s1 = sp.stats.chi2.ppf(q=0.25, df=5)

xL = sp.stats.chi2.ppf(q=0.0001, df=5)
xR = sp.stats.chi2.ppf(q=0.9999, df=5)
x = np.linspace(xL, xR, 1000)
y = sp.stats.chi2.pdf(x=x, df=5)
plt.plot(x, y, 'k')
edge = np.searchsorted(x, s1)
plt.fill_between(x[:edge], y[:edge], color='tab:green')
plt.fill_between(x[edge-2:edge+2], y[edge-2:edge+2], color='red')
plt.title('P(X < {:.4f}) = 0.25'.format(s1))
plt.xlabel('x')
plt.ylabel('y');
```

<!-- grade_id: cell-bc3bb3a5ab319d0b, points: 0, grade: false, solution: true -->
```python
# s2
s2 = sp.stats.f.ppf(q=1-0.15, dfn=5, dfd=20)

xL = sp.stats.f.ppf(q=0.0001, dfn=5, dfd=20)
xR = sp.stats.f.ppf(q=0.9999, dfn=5, dfd=20)
x = np.linspace(xL, xR, 1000)
y = sp.stats.f.pdf(x=x, dfn=5, dfd=20)
plt.plot(x, y, 'k')
edge = np.searchsorted(x, s2)
plt.fill_between(x[edge:], y[edge:], color='tab:green')
plt.fill_between(x[edge-2:edge+2], y[edge-2:edge+2], color='red')
plt.title('P(X > {:.4f}) = 0.15'.format(s2))
plt.xlabel('x')
plt.ylabel('y');
```

<!-- grade_id: cell-835a0cb518f4150e, points: 0, grade: false, solution: true -->
```python
# s3
q = 0.35 + sp.stats.norm.cdf(x=-1.25, loc=-0.4, scale=0.75)
s3 = sp.stats.norm.ppf(q=q, loc=-0.4, scale=0.75)

xL = sp.stats.norm.ppf(q=0.0001, loc=-0.4, scale=0.75)
xR = sp.stats.norm.ppf(q=0.9999, loc=-0.4, scale=0.75)
x = np.linspace(xL, xR, 1000)
y = sp.stats.norm.pdf(x=x, loc=-0.4, scale=0.75)
plt.plot(x, y, 'k')
edge_start = np.searchsorted(x, -1.25)
edge = np.searchsorted(x, s3)
plt.fill_between(x[edge_start:edge], y[edge_start:edge], color='tab:green')
plt.fill_between(x[edge-2:edge+2], y[edge-2:edge+2], color='red')
plt.title('P(-1.25 < X < {:.4f}) = 0.35'.format(s3))
plt.xlabel('x')
plt.ylabel('y');
```

```python
print(s1) # błąd zaokrąglenia
print(s2)
print(s3) # niezidentyfikowany błąd
```

<!-- grade_id: cell-5beafcc4cebdbc21, points: 5, grade: true, solution: false -->
```python
### Komórka testowa
assert s1 == 2.6746028094321637
assert s2 == 1.843305229426929
assert s3 == -0.4835496337042254
```