# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD28"
COLLABORATORS = "Radosław Matuszczyk 405124, Piotr Łukasiewicz 406163"
```

---

## Zestaw zadań 3: Elementy statystyki opisowej 3 (Notatnik 2/4)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import dataframe_image as dfi
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
import matplotlib.patches as mpatches
```

### Wczytanie danych
Wczytaj dane z przygotowanych wcześniej plików bike_traffic_data.pkl i  weather_data.pkl do tablic `bike_traffic_data` oraz `weather_data`.

Wyświetl przygotowane tablice.

<!-- grade_id: cell-184ec13f50fb3bfe, points: 0, grade: true, solution: true -->
```python
# Ruch rowerowy
bike_traffic_data = pd.read_pickle('bike_traffic_data.pkl')
bike_traffic_data
```

<!-- grade_id: cell-cf8ff5e0a8a18cb2, points: 0, grade: true, solution: true -->
```python
# Pogoda
weather_data = pd.read_pickle('weather_data.pkl')
weather_data
```

### Zadnie 2: Przegląd danych - ruch rowerowy [15 pkt]

#### a) Tabele

Przygotuj trzy tablice prezentujące różne aspekty ruchu rowerowego według poniższych wytycznych.

Tablica `daily_bike_traffic_data` prezentuje sumaryczną liczbę rowerzystów przejeżdżających przez most w poszczególnych dniach roku w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu.
> Indeks: Date</br>
> PeriodIndex zawierający okresy czasu, w których dokonywane były zliczenia rowerzystów.</br>

> Kolumny: Total, East, West</br>
> Kolumny zawierają liczbę rowerzystów przejeżdżających przez most w poszczególnych dniach roku w sumie oraz wraz z podziałem na zachodnią i wschodnią stronę mostu.
       
Tablica `bike_traffic_by_weekday_data` prezentuje średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych dniach tygodnia w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu.

> Indeks: Day</br>
> Dni tygodnia (Monday, Tuesday ... ).</br>

> Kolumny: Total, East, West</br>
> Kolumny zawierają średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych dniach tygodnia w sumie oraz wraz z podziałem na zachodnią i wschodnią stronę mostu.
    
Tablica `weekday_weekend_hourly_bike_traffic_data`  prezentuje średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu, a także z podziałem na dni robocze i weekend.

> Indeks: Hour </br>
> Godziny, dla których liczone są wartości średnie (0, 1, 2 ...).

> Kolumny: </br>
> 1 rzędu: Weekday, Weekend</br>
> 2 rzędu: Total, East, West (ten sam zestaw kolumn dla każdego okresu tygodnia)</br>
> Kolumny zawierają średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w sumie oraz wraz z podziałem na zachodnią i wschodnią stronę mostu.

Zapisz tablicę `daily_bike_traffic_data` do pliku daily_bike_traffic_data.pkl.

Wyświetl przygotowane tablice w taki sposób aby zawarte w nich liczby rowerzystów wyświetlały się z dokładnością 0 miejsc po przecinku, a średnie liczby rowerzystów z dokładnością 2 miejsc po przecinku. W przypadku tablicy `daily_bike_traffic_data` wyświetl tylko 10 pierwszych wierszy tablicy.

<!-- grade_id: cell-8be9f2f086031bd0, points: 0, grade: false, solution: true -->
```python
# Agregacja danych z poszczególnych dni roku
daily_bike_traffic_data = bike_traffic_data.resample('D').sum()
daily_bike_traffic_data.head(10)
```

<!-- grade_id: cell-39ea156f9a6713c1, points: 0, grade: false, solution: true -->
```python
# Agregacja danych z poszczególnych dni tygodnia
bike_traffic_by_weekday_data = daily_bike_traffic_data.groupby(daily_bike_traffic_data.index.day_of_week).mean()
bike_traffic_by_weekday_data.index = ['Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday', 'Sunday']
bike_traffic_by_weekday_data.index.name = 'Day'
pd.io.formats.style.Styler(bike_traffic_by_weekday_data, precision=2)
```

<!-- grade_id: cell-fcbd547739059e22, points: 0, grade: false, solution: true -->
```python
# Agregacja danych z konkretnych okresów tygodnia (dzień roboczy/weekend)
weekdays_indexes = bike_traffic_data.index.weekday < 5
weekdays_data = bike_traffic_data[weekdays_indexes].groupby(bike_traffic_data[weekdays_indexes].index.hour).mean()
weekends_data = bike_traffic_data[~weekdays_indexes].groupby(bike_traffic_data[~weekdays_indexes].index.hour).mean()
weekday_weekend_hourly_bike_traffic_data = pd.concat([weekdays_data, weekends_data], axis=1)
weekday_weekend_hourly_bike_traffic_data.index.name = 'Hour'
columns = pd.MultiIndex.from_tuples([('Weekday', 'Total'), 
                                     ('Weekday', 'East'), 
                                     ('Weekday', 'West'), 
                                     ('Weekend', 'Total'), 
                                     ('Weekend', 'East'), 
                                     ('Weekend', 'West')])
weekday_weekend_hourly_bike_traffic_data.columns = columns
pd.io.formats.style.Styler(weekday_weekend_hourly_bike_traffic_data, precision=2)
```

<!-- grade_id: cell-83721cfb90114427, points: 7, grade: true, solution: false -->
```python
### Komórka testowa

## Tabela 1
daily_bike_traffic_data
# Test 1 Kształt tablicy
assert daily_bike_traffic_data.shape == (365, 3)
# Test 2 Kolumny
assert len(set(list(daily_bike_traffic_data.columns)) & set(['Total', 'East', 'West'])) == 3
# Test 3 Indeks
assert (daily_bike_traffic_data.index.name == 'Date') == True
assert isinstance(daily_bike_traffic_data.index, pd.PeriodIndex)
assert daily_bike_traffic_data.index.values.min() == pd.Period('2022-01-01', 'D')
assert daily_bike_traffic_data.index.values.max() == pd.Period('2022-12-31', 'D')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(daily_bike_traffic_data[['Total', 'East', 'West']].min().values, np.array([23., 12., 10.]))) == True
assert np.all(np.isclose(daily_bike_traffic_data[['Total', 'East', 'West']].max().values, np.array([4233., 1733., 3031.]))) == True

## Tabela 2
# Test 1 Kształt tablicy
assert bike_traffic_by_weekday_data.shape == (7, 3)
# Test 2 Kolumny
assert len(set(list(bike_traffic_by_weekday_data.columns)) & set(['Total', 'East', 'West'])) == 3
# Test 3 Indeks
assert (bike_traffic_by_weekday_data.index.name == 'Day') == True
assert len(set(list(bike_traffic_by_weekday_data.index)) & set(['Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday', 'Sunday'])) == 7
# Test 4 Zawartość kolumn
assert np.all(np.isclose(bike_traffic_by_weekday_data.min().values, np.array([1664.44230769,  771.57692308,  892.86538462]))) == True
assert np.all(np.isclose(bike_traffic_by_weekday_data.max().values, np.array([2579.96153846,  983.82692308, 1596.13461538]))) == True

## Tabela 3
# Test 1 Kształt tablicy
assert weekday_weekend_hourly_bike_traffic_data.shape == (24, 6)
# Test 2 Kolumny
assert len(set(list(weekday_weekend_hourly_bike_traffic_data.columns)) & set(list(weekday_weekend_hourly_bike_traffic_data.columns))) == 6
# Test 3 Indeks
assert (weekday_weekend_hourly_bike_traffic_data.index.name == 'Hour') == True
assert weekday_weekend_hourly_bike_traffic_data.index.values.min() == 0
assert weekday_weekend_hourly_bike_traffic_data.index.values.max() == 23
# Test 4 Zawartość kolumn
assert np.all(np.isclose(weekday_weekend_hourly_bike_traffic_data.min().values, np.array([1.75, 0.87307692, 0.87692308, 2.82857143, 1.07619048, 1.53333333]))) == True
assert np.all(np.isclose(weekday_weekend_hourly_bike_traffic_data.max().values, np.array([324.42307692, 99.01923077, 225.72307692, 187.73333333, 87.79047619, 99.94285714]))) == True
```

#### b) Wykresy

Zwizualizuj dane z tablic `daily_bike_traffic_data`, `bike_traffic_by_weekday_data` oraz `weekday_weekend_hourly_bike_traffic_data` za pomocą wykresów liniowych. Przygotuj:
- wykres prezentujący liczbę rowerzystów przejeżdżających przez most w poszczególnych dniach roku w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu,
- wykres prezentujący średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych dniach tygodnia w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu,
- wykres prezentujący średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w dni robocze w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu,
- wykres prezentujący średnią liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w weekendy w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu.

Zestaw wykresy na pojedynczej figurze (siatka wykresów 2x3, przy czym wykres sum dziennych zajmuje wszystkie trzy pola w górnym wierszu figury). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw różne od domyślnych i od siebie nawzajem kolory wykresów poszczególnych parametrów (każdy z parametrów ma na wszystkich wykresach ten sam kolor),
- dodaj do figury pojedynczą legendę zawierającą nazwy zwizualizowanych parametrów i zlokalizowaną w dolnej części figury, pod wykresami,
- dopasuj zakresy osi wykresów do zakresu wartości danych.

<!-- grade_id: cell-202b025e3a1ca82c, points: 8, grade: true, solution: true -->
```python
fig = plt.figure(figsize=(12,8))
axes = []
axes.append(fig.add_subplot(211))
axes.append(fig.add_subplot(234))
axes.append(fig.add_subplot(235))
axes.append(fig.add_subplot(236))
fig.tight_layout(h_pad=7, w_pad=5)
fig.suptitle('Liczba rowerzystów przejeżdzająca przez most', fontsize=16, y=1.05)

colors = ['tab:blue', 'tab:red', 'green']

daily_bike_traffic_data.plot(ax=axes[0], legend=False, color=colors)
axes[0].set_ylim([0,5000])
axes[0].set_title('Suma w poszczególnych dniach')
axes[0].set_ylabel('Sum')

bike_traffic_by_weekday_data.plot(ax=axes[1], legend=False, color=colors)
axes[1].set_xlim([0,6])
axes[1].tick_params(axis='x', labelrotation = 45)
axes[1].set_title('Średnia w poszczególnych dniach tygodnia', fontsize=12)
axes[1].set_ylabel('Mean')

weekday_weekend_hourly_bike_traffic_data['Weekday'].plot(ax=axes[2], legend=False, color=colors)
axes[2].set_xlim([0,23])
axes[2].set_ylim([0,350])
axes[2].set_title('Średnia w poszczególnych godzinach\n(dni robocze)', fontsize=12)
axes[2].set_ylabel('Mean')

weekday_weekend_hourly_bike_traffic_data['Weekend'].plot(ax=axes[3], legend=False, color=colors)
axes[3].set_xlim([0,23])
axes[3].set_ylim([0,350])
axes[3].set_title('Średnia w poszczególnych godzinach\n(weekendy)', fontsize=12)
axes[3].set_ylabel('Mean')

fig.legend(handles=[plt.Line2D([], [], color=colors[0], label='Total'),
                    plt.Line2D([], [], color=colors[1], label='East'),
                    plt.Line2D([], [], color=colors[2], label='West')], loc='lower center', bbox_to_anchor=(0.5, -0.1), ncols=3);

plt.savefig('bikers.png', bbox_inches='tight', pad_inches=0.8)
```

### Zadnie 3: Przegląd danych - pogoda [15 pkt]

#### a) Tabela

Przygotuj tablicę `monthly_weather_data` zawierającą dane z tablicy `weather_data` zagregowane w obrębie poszczególnych miesięcy według następującego klucza:
- wartość minimalna z całego miesiąca: Temperature: Min,
- wartość maksymalna z całego miesiąca: Temperature: Max, Wind: Max speed,
- suma z całego miesiąca: Precipitation: Total, Snow: Fall,
- średnia z całego miesiąca: Temperature: Avg, Snow: Depth, Wind: Avg speed, Sky: Cloud cover.

Tablica powinna zawierać wyłącznie następujące elementy:</br>
> Indeks: Date</br>
> PeriodIndex zawierający okresy czasu, których dotyczą dane z raportów pogodowych.</br>

> Kolumny: W podanej niżej kolejności</br>
> Temperature: Min, Avg, Max</br>
> Precipitation: Total</br>
> Snow: Fall, Depth</br>
> Wind: Avg speed, Max speed</br>
> Sky: Cloud cover</br>
> (Nazwa 1 poziomu: podległe jej nazwy 2 poziomu)

Wyświetl przygotowaną tablicę w taki sposób aby zawarte w nich wartości wyświetlały się z dokładnością 2 miejsc po przecinku.

<!-- grade_id: cell-6d40159c1f916fef, points: 0, grade: false, solution: true -->
```python
functions = ['min', 'mean', 'max', 'sum', 'sum', 'mean', 'mean', 'max', 'mean']
dictMap =  dict(zip(weather_data.columns, functions))

monthly_weather_data = weather_data.resample('M').sum()
for key, value in dictMap.items():
    if value == 'mean':
        monthly_weather_data[key] = weather_data[key].resample('M').sum() / [31, 28, 31, 30, 31, 30, 31, 31, 30, 31, 30, 31]
    else:
        monthly_weather_data[key] = weather_data[key].resample('M').agg(value)

pd.io.formats.style.Styler(monthly_weather_data, precision=2)
```

<!-- grade_id: cell-c2c51d232639ab46, points: 7, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert monthly_weather_data.shape == (12, 9)
# Test 2 Kolumny
assert len(set(list(monthly_weather_data.columns)) & set([('Temperature', 'Min'), ('Temperature', 'Avg'), ('Temperature', 'Max'), ('Precipitation', 'Total'),
            ('Snow', 'Fall'), ('Snow', 'Depth'), ('Wind', 'Avg speed'), ('Wind', 'Max speed'), ('Sky', 'Cloud cover')])) == 9
# Test 3 Indeks
assert (monthly_weather_data.index.name == 'Date') == True
assert isinstance(monthly_weather_data.index, pd.PeriodIndex) == True
assert monthly_weather_data.index.values.min() == pd.Period('2022-01', 'M')
assert monthly_weather_data.index.values.max() == pd.Period('2022-12', 'M')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(monthly_weather_data.min().values, np.array([-7.77777778, 3.90681004, 13.33333333, 1.27, 0., 0., 2.93603606, 8.0467, 0.43870968]))) == True
assert np.all(np.isclose(monthly_weather_data.max().values, np.array([11.66666667, 21.27240143, 35., 191.77, 119.38, 4.91612903, 3.97864611, 16.54043889, 0.86774194]))) == True
```

#### b) Wykresy

Zwizualizuj dane z tablicy `weather_data` za pomocą następującego zestawu wykresów:
- temperatura dzienna (minimalna, średnia i maksymalna), wypełnij kolorem przestrzeń między temperaturą minimalną i maksymalną,
- dzienne zachmurzenie (lewa oś wykresu) i opady (prawa oś wykresu), wypełnij kolorem przestrzeń pod każdym z wykresów,
- dzienna średnia i maksymalna prędkość wiatru.

Zestaw wykresy na pojedynczej figurze (siatka wykresów 3x1). Przygotuj figurę zgodnie z wytycznymi z nagłówka notatnika. Dodatkowo:
- ustaw różne od domyślnych i od siebie nawzajem kolory wykresów poszczególnych parametrów,
- dodaj do figury pojedynczą legendę zawierającą nazwy wizualizowanych parametrów i zlokalizowaną w dolnej części figury, pod wykresami, uwzględnij w legendzie nie tylko kolory linii, ale również zastosowanych wypełnień,
- dopasuj zakresy osi wykresów do zakresu danych, w szczególności dobierz skale osi Y tak, aby wykresy opadów i zachmurzenia nie przysłaniały się wzajemnie,
- w opisach osi Y umieść jednostki, w których wyrażone są wartości widoczne na wykresach.

<!-- grade_id: cell-b72d17b7c9b43a58, points: 8, grade: true, solution: true -->
```python
fig, axes = plt.subplots(3, 1, figsize=(12,10))
fig.tight_layout(h_pad=6)
fig.suptitle('Pogoda w roku 2022', fontsize=16, y=1.05)

colors = ['tab:blue', 'green', 'tab:red', 'lightgreen', 'dimgrey', 'silver', 'blue', 'skyblue', 'gold', 'orange']

weather_data['Temperature'].plot(ax=axes[0], legend=False, color=colors[:3])
axes[0].fill_between(weather_data.index, weather_data['Temperature']['Max'], weather_data['Temperature']['Min'],
                  alpha=0.5, color=colors[3])
axes[0].set_ylim([-10,40])
axes[0].set_title('Temperatura')
axes[0].set_ylabel('Temperatura [°C]')

weather_data['Sky'].plot(ax=axes[1], legend=False, color=colors[4])
axes[1].fill_between(weather_data.index, weather_data['Sky']['Cloud cover'], alpha=0.5, color=colors[5])
axes[1].set_ylim([0,1])
axes[1].set_title('Zachmurzenie i opady')
axes[1].set_ylabel('Zachmurznenie [ułamek]', color=colors[4])
axes[1].tick_params(axis ='y', labelcolor=colors[4])
axR1 = axes[1].twinx()
weather_data['Precipitation'].plot(ax=axR1, legend=False, color=colors[6])
axR1.fill_between(weather_data.index, weather_data['Precipitation']['Total'], alpha=0.5, color=colors[7])
axR1.set_ylabel('Opady [mm]', color=colors[6])
axR1.tick_params(axis ='y', labelcolor=colors[6])
axR1.set_ylim([0,80])

weather_data['Wind'].plot(ax=axes[2], legend=False, color=colors[8:])
axes[2].set_ylim([0,20])
axes[2].set_title('Wiatr')
axes[2].set_ylabel('Prędkość wiatru [m/s]')

fig.legend(handles=[plt.Line2D([], [], color=colors[0], label='Temperature, Max'),
                    plt.Line2D([], [], color=colors[1], label='Temperature, Avg'),
                    plt.Line2D([], [], color=colors[2], label='Temperature, Min'),
                    mpatches.Rectangle((0,0),1,1, color=colors[3], label='Temperature, Range'),
                    mpatches.Rectangle((0,0),1,1, edgecolor=colors[4], facecolor=colors[5], label='Temperature, Range'),
                    mpatches.Rectangle((0,0),1,1, edgecolor=colors[6], facecolor=colors[7], label='Temperature, Range'),
                    plt.Line2D([], [], color=colors[8], label='Wind, Avg speed'),
                    plt.Line2D([], [], color=colors[9], label='Wind, Max speed')], 
           loc='lower center', bbox_to_anchor=(0.5, -0.075), ncols=8, fontsize=8);
```