# Analiza danych w języku Python - ćwiczenia laboratoryjne 2022/2023

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (`NAME` - nazwa zespołu, `COLLABORATORS` - imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

```python
NAME = "IAD28"
COLLABORATORS = "Radosław Matuszczyk 405124, Piotr Łukasiewicz 406163"
```

---

## Zestaw zadań 3: Elementy statystyki opisowej 3 (Notatnik 1/4)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
import dataframe_image as dfi
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Przygotowanie danych [10 pkt]

Pliki bike_traffic_raw_data.csv, 2022_01.txt, 2022_02.txt ...  2022_12.txt zawierają dane pobrane ze stron:</br>
https://data.seattle.gov/Transportation/Fremont-Bridge-Bicycle-Counter/65db-xm6k,</br>
https://forecast.weather.gov/product.php?site=SEW&product=CF6&issuedby=SEW.

#### a) Natężenie ruchu rowerowego 

Plik raw_bike_traffic_data.csv zawiera dane na temat ruchu rowerowego na moście Fremont Bridge w Seattle.

Zaimportuj zawarte w pliku dane do notatnika i utwórz na ich podstawie tablicę `bike_traffic_data` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    > Indeks: Date</br>
    > PeriodIndex zawierający okresy czasu, w których dokonywane były zliczenia rowerzystów.
    
    > Kolumny: Total, East, West</br>
    > Kolumny zawierają liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu.
- Pozostaw w tablicy wyłącznie wiersze dotyczące 2022 roku.

Zapisz tablicę do pliku bike_traffic_data.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 12 pierwszych i ostatnich wierszy, a zawarte w tablicy wartości wyświetlały się bez miejsc po przecinku.

<!-- grade_id: cell-aa85199b5cb633d2, points: 0, grade: false, solution: true -->
```python
bike_traffic_data = pd.read_csv('raw_bike_traffic_data.csv', index_col=0)
bike_traffic_data.columns = ['Total', 'East', 'West'] 
bike_traffic_data.index = pd.to_datetime(bike_traffic_data.index).to_period()
bike_traffic_data = bike_traffic_data[bike_traffic_data.index.year == 2022]
bike_traffic_data = bike_traffic_data.convert_dtypes()

bike_traffic_data.to_pickle('bike_traffic_data.pkl')

pd.io.formats.style.Styler(bike_traffic_data).hide(subset=bike_traffic_data.index[12:-12])
```

Output:
```
<pandas.io.formats.style.Styler at 0x217d60d7010>
```

<!-- grade_id: cell-683b0a3c6c74d0c6, points: 2, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert bike_traffic_data.shape == (8760, 3)
# Test 2 Kolumny
assert len(set(list(bike_traffic_data.columns)) & set(['Total', 'East', 'West'])) == 3
# Test 3 Indeks
assert (bike_traffic_data.index.name == 'Date') == True
assert isinstance(bike_traffic_data.index, pd.PeriodIndex)
assert bike_traffic_data.index.values.min() == pd.Period('2022-01-01 00:00', 'H')
assert bike_traffic_data.index.values.max() == pd.Period('2022-12-31 23:00', 'H')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(bike_traffic_data[['Total', 'East', 'West']].min().values, np.array([0, 0, 0]))) == True
assert np.all(np.isclose(bike_traffic_data[['Total', 'East', 'West']].max().values, np.array([717, 464, 635]))) == True
```

#### b) Pogoda

Pliki 2022_01.txt, 2022_02.txt ...  2022_12.txt zawierają miesięczne raporty pogodowe (F6 Preliminary Local Climate Data Report) ze stacji NWS Seattle/Tacoma, WA.
Wyjaśnienie wykorzystanych w raportach symboli i skrótów można znaleźć pod następującym linkiem: https://www.weather.gov/grr/climateF6explain.

Zaimportuj zawarte w plikach dane do notatnika i utwórz na ich podstawie tablicę `weather_data` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    > Indeks: Date</br>
    > PeriodIndex zawierający okresy czasu, których dotyczą dane z raportów pogodowych.</br>
    
    > Kolumny: W podanej niżej kolejności</br>
    > Temperature: Min, Avg, Max</br>
    > Precipitation: Total</br>
    > Snow: Fall, Depth</br>
    > Wind:  Avg speed, Max speed</br>
    > Sky: Cloud cover</br>
    > (Nazwa 1 poziomu: podległe jej nazwy 2 poziomu)
- Zamień oznaczenia opadów śladowych na wartość 0.
- Przelicz występujące w tablicy wartości do *C (temperatura), mm (opady), m/s (wiatr) i ułamków (zachmurzenie).

Zapisz tablicę do pliku weather_data.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 10 pierwszych i ostatnich wierszy, a zawarte w tablicy liczby wyświetlały się z dokładnością 2 miejsc po przecinku.

<!-- grade_id: cell-186c2f35c4f9aca6, points: 0, grade: false, solution: true -->
```python
weather_data = pd.DataFrame()
daysInMonths = [31, 28, 31, 30, 31, 30, 31, 31, 30, 31, 30, 31]
for ix, daysInMonth in enumerate(daysInMonths):
    df = pd.read_table('raw_weather_data/2022_{:02}.txt'.format(ix+1), skiprows=15, nrows=daysInMonth+1, sep='\s+', na_values='M').drop(0)
    weather_data = pd.concat((weather_data, df))

start_date = '2022-01-01'
end_date = '2022-12-31'
weather_data.index = period_index = pd.period_range(start=start_date, end=end_date, freq='D')
weather_data.index.name = 'Date'

columns = ['MIN', 'AVG', 'MAX', 'WTR', 'SNW', 'DPTH', 'SPD', 'SPD.1', 'S-S']
weather_data = weather_data[columns]

levels = [('Temperature', 'Min'), ('Temperature', 'Avg'), ('Temperature', 'Max'),
          ('Precipitation', 'Total'),
          ('Snow', 'Fall'), ('Snow', 'Depth'),
          ('Wind', 'Avg speed'), ('Wind', 'Max speed'),
          ('Sky', 'Cloud cover')] 
weather_data.columns = pd.MultiIndex.from_tuples(levels)

weather_data.replace('T', 0, inplace=True)
weather_data = weather_data.astype('float')
weather_data['Temperature'] = (weather_data['Temperature']-32)/1.8
weather_data['Precipitation'] = weather_data['Precipitation']*25.4
weather_data['Snow'] = weather_data['Snow']*25.4
weather_data['Wind'] = weather_data['Wind']*0.44704
weather_data['Sky'] = weather_data['Sky']/10

weather_data.to_pickle('weather_data.pkl')

pd.io.formats.style.Styler(weather_data, precision=2).hide(subset=weather_data.index[10:-10])
```

Output:
```
<pandas.io.formats.style.Styler at 0x217e766c940>
```

<!-- grade_id: cell-e70d426962ca7e21, points: 8, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert weather_data.shape == (365, 9)
# Test 2 Kolumny
assert len(set(list(weather_data.columns)) & set([('Temperature', 'Min'), ('Temperature', 'Avg'), ('Temperature', 'Max'), ('Precipitation', 'Total'),
            ('Snow', 'Fall'), ('Snow', 'Depth'), ('Wind', 'Avg speed'), ('Wind', 'Max speed'), ('Sky', 'Cloud cover')])) == 9
# Test 3 Indeks
assert (weather_data.index.name == 'Date') == True
assert isinstance(weather_data.index, pd.PeriodIndex) == True
assert weather_data.index.values.min() == pd.Period('2022-01-01', 'D')
assert weather_data.index.values.max() == pd.Period('2022-12-31', 'D')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(weather_data.min().values, np.array([-7.77777778, -5.55555556, -3.33333333, 0., 0., 0., 0.89407778, 3.12927222, 0.]))) == True
assert np.all(np.isclose(weather_data.max().values, np.array([20., 26.66666667, 35., 75.438, 66.04, 76.2, 8.31492333, 16.54043889, 1.]))) == True
```